# SLAM Unit A - Part 1
In this first unit, Unit A, we will learn how to use the motor and the scanner.
This is the first part of the first unit, in which we will get acquainted with the notebook and perform first experiments with the robots motor.

## Introduction
Here, our small robot and the overall setup of our experiment are introduced.

```python
# A technical note: uncomment the following line to get YouTube videos instead of TIB AV.
# YouTube = True
# Note: to go back to TIB AV, you will have to comment the line out again and do a Kernel --> Restart.
```

```python
# If you don't see a video below, run this cell.
from IPython.display import IFrame
IFrame("https://www.youtube.com/embed/sxu6_YnZca8" if "YouTube" in globals() else "//av.tib.eu/player/48877",
       width=560, height=315)
```

### Here is our first question (1 Point):
What is the number of landmarks in our arena?

<!-- grade_id: cell-d56f6868db087500, points: 0, grade: false, solution: true -->
```python
the_number_of_landmarks = 6  # Replace this by the correct number and delete any lines below.
# YOUR CODE HERE
# raise NotImplementedError()
```

<!-- grade_id: cell-ead953104637316e, points: 1, grade: true, solution: false -->
```python
# This notebook will test your answer using public and hidden tests.
# The public tests are meant to help you to develop your solution.
# The hidden tests will check your solution for the grading.

# The following lines are the public tests.
# In this case, the tests will not make much sense, but they demonstrate how they work in general.

# It does not make sense to have a negative number of landmarks, and also we will have at least one.
assert(the_number_of_landmarks > 0)

# Also, since our arena is quite small, we would be unable to place 100 or more landmarks.
assert(the_number_of_landmarks < 100)

```

## Motor control

```python
# If you don't see a video below, run this cell.
from IPython.display import IFrame
IFrame("https://www.youtube.com/embed/Wt-ivnBZsv0" if "YouTube" in globals() else "//av.tib.eu/player/49050",
       width=560, height=315)
```

### After viewing the video, is is easy to answer the following question (1 Point):
What is the name of the Python class that is used to read in the log files of our robot?

<!-- grade_id: cell-3a216bdb78b3c420, points: 0, grade: false, solution: true -->
```python
# Write here the name of the class, keeping the quotes around it. Then, delete any lines below.
the_name_of_the_class = "import"
# YOUR CODE HERE
# raise NotImplementedError()
```

<!-- grade_id: cell-b7605756340906af, points: 1, grade: true, solution: false -->
```python
# Now, some tests.
assert(the_name_of_the_class != None)  # That would not be a valid answer.
assert(the_name_of_the_class != "")  # That one neither.
assert(the_name_of_the_class != "I don't know")  # Well, check the video...

```

### In addition, let us redo the steps from the video. First, we plot the motor outputs.
You may play with the code below and re-execute it in order to get different visualizations.
If you are curious about how plotting works, check out pylab, which is a part of __[matplotlib](https://matplotlib.org/)__.

```python
# You can modify the following to adjust the (width, height) of the plots (supposed to be in inches).
from pylab import *
plt.rcParams['figure.figsize'] = (12, 6)
```

```python
# Plot the ticks from the left and right motor.
# 01_a_plot_motor_ticks.py
from pylab import *

if __name__ == '__main__':
    # Read all ticks of left and right motor.
    # Format is:
    # M timestamp[in ms] pos-left[in ticks] * * * pos-right[in ticks] ...
    # so we are interested in field 2 (left) and 6 (right).
    f = open("robot4_motors.txt")
    left_list = []
    right_list = []
    for l in f:
        sp = l.split()
        left_list.append(int(sp[2]))
        right_list.append(int(sp[6]))

    plot(left_list, c="blue")
    plot(right_list, c="green")
    show()
```

### Then, we read the ticks using the provided class and output the first few of them.
Now, these are not the absolute values, but the increments.

```python
# Print the increments of the left and right motor.
# Now using the LegoLogfile class.
# 01_b_print_motor_increments.py
from lego_robot import LegoLogfile

if __name__ == '__main__':

    logfile = LegoLogfile()
    logfile.read("robot4_motors.txt")

    for i in range(20):
        print(logfile.motor_ticks[i])
```

### Finally, we plot the increments.

```python
# Plot the increments of the left and right motor.
# 01_c_plot_motor_increments.py
from pylab import *
from lego_robot import LegoLogfile

if __name__ == '__main__':

    logfile = LegoLogfile()
    logfile.read("robot4_motors.txt")

    # Actually this statement does it: plot(logfile.motor_ticks), but in order to
    # reproduce the colors in the video, I will make it slightly more complicated:
    plot([t[0] for t in logfile.motor_ticks], c="blue")
    plot([t[1] for t in logfile.motor_ticks], c="green")
    show()
```