# SLAM Unit D - Part 2
Introducing the Kalman filter for our robot.

```python
# If you don't see a video below, run this cell.
# YouTube = True  # Uncomment to get YouTube videos instead of TIB AV.
from IPython.display import IFrame
IFrame("https://www.youtube.com/embed/cXx5NmtPFag" if "YouTube" in globals() else "//av.tib.eu/player/49004",
       width=560, height=315)
```

Output:
```
<IPython.lib.display.IFrame at 0x7f653c248910>
```

<!-- grade_id: cell-085e9943175910dd, points: 0, grade: false, solution: true -->
```python
# What is the dimension of A? Write this as a string, e.g. "1x2" for "one row times 2 columns".
dimension_of_A = "3x3"
# YOUR CODE HERE
#raise NotImplementedError()
```

<!-- grade_id: cell-b2767ca1d836c40c, points: 1, grade: true, solution: false -->
```python
from hashlib import shake_128
def public_string_test(the_answer_string, reference):
    m = shake_128()
    m.update(the_answer_string.encode())
    return m.hexdigest(4) == reference
assert public_string_test("A="+dimension_of_A.strip().lower(), '6bdfd84e'), "Oh no, your answer is wrong!"
```

```python
# If you don't see a video below, run this cell.
from IPython.display import IFrame
IFrame("https://www.youtube.com/embed/VlvaZtABDoo" if "YouTube" in globals() else "//av.tib.eu/player/49005",
       width=560, height=315)
```

Output:
```
<IPython.lib.display.IFrame at 0x7f653c2b7290>
```

<!-- grade_id: cell-635335df1bc9e6a4, points: 0, grade: false, solution: true -->
```python
# What is the dimension of B? Write this as a string, e.g. "1x2" for "one row times 2 columns".
dimension_of_B = "3x2"
# YOUR CODE HERE
#raise NotImplementedError()
```

<!-- grade_id: cell-4318d3278d4a4a71, points: 1, grade: true, solution: false -->
```python
from hashlib import shake_128
def public_string_test(the_answer_string, reference):
    m = shake_128()
    m.update(the_answer_string.encode())
    return m.hexdigest(4) == reference
assert public_string_test("B="+dimension_of_B.strip().lower(), 'eaa7d686'), "Oh no, your answer is wrong!"
```

```python
# If you don't see a video below, run this cell.
from IPython.display import IFrame
IFrame("https://www.youtube.com/embed/lMLKk7knWxk" if "YouTube" in globals() else "//av.tib.eu/player/49006",
       width=560, height=315)
```

Output:
```
<IPython.lib.display.IFrame at 0x7f5f744da550>
```

<!-- grade_id: cell-cedfe497ba2146a3, points: 0, grade: false, solution: true -->
```python
# What is the dimension of C? Write this as a string, e.g. "1x2" for "one row times 2 columns".
dimension_of_C = "2x3"
# YOUR CODE HERE
#raise NotImplementedError()
```

<!-- grade_id: cell-53c690fd5eb4c87a, points: 1, grade: true, solution: false -->
```python
from hashlib import shake_128
def public_string_test(the_answer_string, reference):
    m = shake_128()
    m.update(the_answer_string.encode())
    return m.hexdigest(4) == reference
assert public_string_test("C="+dimension_of_C.strip().lower(), '3acb778a'), "Oh no, your answer is wrong!"
```

```python
# If you don't see a video below, run this cell.
from IPython.display import IFrame
IFrame("https://www.youtube.com/embed/g6qqfyN-7Ec" if "YouTube" in globals() else "//av.tib.eu/player/49007",
       width=560, height=315)
```

Output:
```
<IPython.lib.display.IFrame at 0x7f653c2b7bd0>
```

<!-- grade_id: cell-91a949be84e42c18, points: 0, grade: false, solution: true -->
```python
# What is the dimension of R? Write this as a string, e.g. "1x2" for "one row times 2 columns".
dimension_of_R = "3x3"
# YOUR CODE HERE
#raise NotImplementedError()
```

<!-- grade_id: cell-990020f117f931b5, points: 1, grade: true, solution: false -->
```python
from hashlib import shake_128
def public_string_test(the_answer_string, reference):
    m = shake_128()
    m.update(the_answer_string.encode())
    return m.hexdigest(4) == reference
assert public_string_test("R="+dimension_of_R.strip().lower(), 'fdbe5ee1'), "Oh no, your answer is wrong!"
```

```python
# If you don't see a video below, run this cell.
from IPython.display import IFrame
IFrame("https://www.youtube.com/embed/J0WqOvCSqA0" if "YouTube" in globals() else "//av.tib.eu/player/49008",
       width=560, height=315)
```

Output:
```
<IPython.lib.display.IFrame at 0x7f653c2b7fd0>
```

<!-- grade_id: cell-0872e943b8b94930, points: 0, grade: false, solution: true -->
```python
# What is the dimension of Q? Write this as a string, e.g. "1x2" for "one row times 2 columns".
dimension_of_Q = "2x2"
# YOUR CODE HERE
#raise NotImplementedError()
```

<!-- grade_id: cell-341b488f41e403b2, points: 1, grade: true, solution: false -->
```python
from hashlib import shake_128
def public_string_test(the_answer_string, reference):
    m = shake_128()
    m.update(the_answer_string.encode())
    return m.hexdigest(4) == reference
assert public_string_test("Q="+dimension_of_Q.strip().lower(), 'e80a899c'), "Oh no, your answer is wrong!"
```

```python
# If you don't see a video below, run this cell.
from IPython.display import IFrame
IFrame("https://www.youtube.com/embed/gdSEp58hOeo" if "YouTube" in globals() else "//av.tib.eu/player/49009",
       width=560, height=315)
```

Output:
```
<IPython.lib.display.IFrame at 0x7f5f76633b50>
```

<!-- grade_id: cell-4d384394d0289850, points: 0, grade: false, solution: true -->
```python
# What is the dimension of sigma? Write this as a string, e.g. "1x2" for "one row times 2 columns".
dimension_of_sigma = "3x3"
# YOUR CODE HERE
#raise NotImplementedError()
```

<!-- grade_id: cell-603ae65ce6ce6090, points: 1, grade: true, solution: false -->
```python
from hashlib import shake_128
def public_string_test(the_answer_string, reference):
    m = shake_128()
    m.update(the_answer_string.encode())
    return m.hexdigest(4) == reference
assert public_string_test("sigma="+dimension_of_sigma.strip().lower(), '51b2f844'), "Oh no, your answer is wrong!"
```

```python
# If you don't see a video below, run this cell.
from IPython.display import IFrame
IFrame("https://www.youtube.com/embed/QBcdaPfCXW8" if "YouTube" in globals() else "//av.tib.eu/player/49010",
       width=560, height=315)
```

Output:
```
<IPython.lib.display.IFrame at 0x7f653c248490>
```

<!-- grade_id: cell-0ab1d947df889e25, points: 0, grade: false, solution: true -->
```python
# What is the dimension of K? Write this as a string, e.g. "1x2" for "one row times 2 columns".
dimension_of_K = "1x2"
# YOUR CODE HERE
raise NotImplementedError()
```

<!-- grade_id: cell-2a77a9fdd1195e2a, points: 1, grade: true, solution: false -->
```python
from hashlib import shake_128
def public_string_test(the_answer_string, reference):
    m = shake_128()
    m.update(the_answer_string.encode())
    return m.hexdigest(4)== reference
assert public_string_test("K="+dimension_of_K.strip().lower(), '6854ccec'), "Oh no, your answer is wrong!"
```