# SLAM Unit D - Part 1
In this unit, you will learn how to apply the Kalman filter to our robot. See the next video for a recapitulation of some previous results and a summary of what you will learn in this unit.

First, however, we will have a look at normal distributions and error ellipses.

```python
# If you don't see a video below, run this cell.
# YouTube = True  # Uncomment to get YouTube videos instead of TIB AV.
from IPython.display import IFrame
IFrame("https://www.youtube.com/embed/xN6OEwudmwo" if "YouTube" in globals() else "//av.tib.eu/player/48999",
       width=560, height=315)
```

Output:
```
<IPython.lib.display.IFrame at 0x7fb5483cff10>
```

## Multivariate Normal Distributions and Error Ellipses

```python
# If you don't see a video below, run this cell.
from IPython.display import IFrame
IFrame("https://www.youtube.com/embed/tP3UAE8dyuw" if "YouTube" in globals() else "//av.tib.eu/player/49000",
       width=560, height=315)
```

<!-- grade_id: cell-ac5fa4901cda3615, points: 0, grade: false, solution: true -->
```python
# Assign each variable a value of False or True, depending on your answer to the question in the video.
A = True  # Or set this to True  (for the circle)
B = False  # Or set this to True  (for the axis-parallel ellipse)
C = True  # Or set this to True  (for the tilted ellipse)
# YOUR CODE HERE
#raise NotImplementedError()
```

<!-- grade_id: cell-7435862ee6712bc7, points: 4, grade: true, solution: false -->
```python
from hashlib import shake_128
def public_string_test(the_answer_string, reference):
    m = shake_128()
    m.update(the_answer_string.encode())
    return m.hexdigest(4) == reference
assert public_string_test(str(bool(A))+str(bool(B))+str(bool(C)), '0de72f87'), "Oh no, your answer is wrong!"
```

### Quiz: the lever arm.

```python
# If you don't see a video below, run this cell.
from IPython.display import IFrame
IFrame("https://www.youtube.com/embed/2KhE8PDqj7M" if "YouTube" in globals() else "//av.tib.eu/player/49001",
       width=560, height=315)
```

Output:
```
<IPython.lib.display.IFrame at 0x7fb54427b250>
```

<!-- grade_id: cell-ec1db75db908ce3e, points: 0, grade: false, solution: true -->
```python
# Is it the shape as shown in A, B, C, or D? Set the string accordingly.
ellipse_is_shaped_like = "D"
# YOUR CODE HERE
#raise NotImplementedError()
```

<!-- grade_id: cell-fbc94ef86329e611, points: 4, grade: true, solution: false -->
```python
from hashlib import shake_128
def public_string_test(the_answer_string, reference):
    m = shake_128()
    m.update(the_answer_string.encode())
    return m.hexdigest(4) == reference
assert public_string_test(ellipse_is_shaped_like.strip().lower(), '5943bb0f'), "Oh no, your answer is wrong!"
```

```python
# If you don't see a video below, run this cell.
from IPython.display import IFrame
IFrame("https://www.youtube.com/embed/cMAoOxMnCsk" if "YouTube" in globals() else "//av.tib.eu/player/49002",
       width=560, height=315)
```

Output:
```
<IPython.lib.display.IFrame at 0x7fb5442a2b90>
```

<!-- grade_id: cell-19f78584d5484efb, points: 0, grade: false, solution: true -->
```python
# What is the correct solution, A, B, C, or D? Set the string accordingly.
variance_is_like_in_case = "C"
# YOUR CODE HERE
#raise NotImplementedError()
```

<!-- grade_id: cell-4cbe462a2bd9cc0f, points: 2, grade: true, solution: false -->
```python
from hashlib import shake_128
def public_string_test(the_answer_string, reference):
    m = shake_128()
    m.update(the_answer_string.encode())
    return m.hexdigest(4) == reference
assert public_string_test(variance_is_like_in_case.strip().lower(), 'a5c1009b'), "Oh no, your answer is wrong!"
```

```python
# If you don't see a video below, run this cell.
from IPython.display import IFrame
IFrame("https://www.youtube.com/embed/9IZMrNNg6Ac" if "YouTube" in globals() else "//av.tib.eu/player/49003",
       width=560, height=315)
```

Output:
```
<IPython.lib.display.IFrame at 0x7fb5442839d0>
```

<!-- grade_id: cell-21f5f5ebccbec74b, points: 0, grade: false, solution: true -->
```python
# What is the correct solution, A, B, C, or D? Set the string accordingly.
half_axis_is_like_in_case = "D"
# YOUR CODE HERE
#raise NotImplementedError()
```

<!-- grade_id: cell-f3e7ace0ceda6889, points: 2, grade: true, solution: false -->
```python
from hashlib import shake_128
def public_string_test(the_answer_string, reference):
    m = shake_128()
    m.update(the_answer_string.encode())
    return m.hexdigest(4) == reference
assert public_string_test(half_axis_is_like_in_case.strip().upper(), 'f0695584'), "Oh no, your answer is wrong!"
```

...join me in part 2 to see why!